Before you turn this problem in, make sure everything runs as expected. First, **Restart and Run All** (in the menubar in colab, select Runtime$\rightarrow$Restart and Run all) 

Make sure you only fill in any place that says `YOUR CODE HERE` and do not make any other changes to the code. If you add any code cells for your own testing. Please delete them before submission.

---

## 1
Write a function that takes a dictionary ```d``` and value ```n``` and returns a list of ```n``` random terms. <br>
Key-value pairs in dictionary ```d``` will be term-probability pairs defining a discrete random variable. Sum of all values in ```d``` will be 1. The function you write should generate ```n``` terms randomly (with replacement) according to the probability distribution given in the dictionary. <br>
For eg., if term ```hello``` has probability ```0.1```, if we generate 1000 terms using the function, the expected number of times ```hello``` occurs should be ```100```

<!-- grade_id: cell-f04c2690de9e5197, points: 0, grade: false, solution: true -->
```python
import numpy as np
def random_generator(d, n):
  """
  Inputs:
    d : dictionary of term-probability pairs
    n : int, number of random terms to generate
  Outputs:
    rand_list : list of length n, contains random terms generated according to probability distribution given in dictionary d
  """
  
  # YOUR CODE HERE
  rand_list = []
  key = []
  value = []
  for i, j in d.items():
    key.append(i)
    value.append(j)

  for k in range(n):
    probability = np.random.choice(key, p = value)
    rand_list.append(probability)
  
  return rand_list
```

<!-- grade_id: cell-650351e8916fecef, points: 0, grade: true, solution: false -->
```python
from collections import Counter
import numpy as np
d = {'hello': 0.5, 'world': 0.5}
rand_lists = []
for _ in range(100):
  rand_lists += random_generator(d, 10)
c = Counter(rand_lists)
assert np.isclose(c['hello'], 500, atol = 100) ## Number of hellos should ideally be close to 500, but we are checking that it should be between 400 and 600
assert np.isclose(c['world'], 500, atol = 100)

```

## 2
Given a numpy array that may contain some nan values, return sum of elements of the array after replacing all ```nan``` values with value ```1```.

<!-- grade_id: cell-7d1764292e628bc4, points: 0, grade: false, solution: true -->
```python
def nan_sum(a):
  """
  Inputs:
  a : numpy array of arbitrary shape

  Outputs:
  sum_ : float, sum of elements of a after replacing nans with 1's
  """
  
  # YOUR CODE HERE
  a[np.isnan(a)]=1.0
  sum_=np.sum(a)
  return sum_
```

<!-- grade_id: cell-6eb7b59b7bbffcf5, points: 0, grade: true, solution: false -->
```python
a = np.ones((3, 3))*2
a[0, 0] = np.nan
a[0, 2] = np.nan
assert np.isclose(nan_sum(a), 16) 

```

## 3
Write a function ```gradient_descent``` to perform gradient descent. <br>
The function takes arguments - ```f```, ```df```, ```startx```, ```lr``` and ```steps```. <br>
```f``` is the function we are trying to minimize. <br>
```df``` is the derivative function of ```f``` <br>
```startx``` is the point ```x``` at which we begin gradient descent <br>
```lr``` is learning rate that we use <br>
```steps``` is number of steps for which we will perform gradient descent


<!-- grade_id: cell-75ff3a85e90b3df3, points: 0, grade: false, solution: true -->
```python
def gradient_descent(f, df, startx, lr, steps):
  """
  Inputs:
    f : function, single argument function which takes a float and returns a float, function to be minimized 
    df : function single argument function which takes a float and returns a float, derivative of f
    startx : float, 
    lr : float,
    steps: int
  Outputs:
    flist : list of length steps+1, values of f at each step of gradient descent 
  """
  
  # YOUR CODE HERE
  flist=[]
  flist.append (f(startx))
  for i in range(steps):
        startx-=lr*df(startx)
        flist.append (f(startx))
  
  return flist
```

<!-- grade_id: cell-0de435e711b9b89d, points: 0, grade: true, solution: false -->
```python
f = lambda x: x**2 - 2*x + 3
df = lambda x: 2*x - 2
startx = 10
lr = 0.2
steps = 100
flist = gradient_descent(f, df, startx, lr, steps)
assert flist[50] <= flist[0]

```

## 4
### Unrolling an RNN
We define a recurrent neuron as follows - 
$$
  y_{t+1} = \begin{cases}
  2*y_t + (y_{t-1}\bmod 3) \text{  if y_t is even} \\
  3*y_t + 1 \text{ if y_t is odd} 
  \end{cases}
$$

Write the function custom_rnn which takes as inputs ```y0``` and ```steps``` and unrolls the RNN defined above for ```steps``` number of steps. (Assume $y_{t-1}$ is 0 for all values of t less than 0)

<!-- grade_id: cell-ae63f5449d761d1c, points: 0, grade: false, solution: true -->
```python
def custom_rnn(y0, steps):
  """
  Inputs:
    y0 : float
    steps: int, number of steps
  Outputs:
    ys : numpy array, array of length steps with values of y
  """
  
  # YOUR CODE HERE
  ys=[]
  ys.append(y0)
  for i in range(steps-1):
        if (ys[i]%2==0):
            y=2*ys[i]
            if (i-1>=0):
                y+=ys[i-1]%3
            ys.append(y)
        else:
            ys.append(3*ys[i]+1)

  ys=np.array(ys,dtype=float)
  return ys
```

<!-- grade_id: cell-4df2a479e9dbeebe, points: 0, grade: true, solution: false -->
```python
ys = custom_rnn(3, 5)
assert np.all(np.isclose(ys, np.array([  3.,  10.,  20.,  41., 124.])))

```

```python

```